<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="category">
    <category><text>$course$/Mastering Ethereum/Section 5 - Wallets (Exercises)</text></category>
  </question>

  <question type="coderunner">
    <name><text>S5 Q1 — Build a BIP-44 path</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S5 Q1 — Derivation paths</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">BIP-44 fixes the shape <code>m / 44' / coin_type' / account' / change / address_index</code>.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write <strong><code>bip44_path(coin_type, account, change, index)</code></strong> that returns the full path string. The purpose level is always <code>44'</code>; the first three numeric levels are hardened (a trailing apostrophe), the last two are not. For example <code>bip44_path(60, 0, 0, 2)</code> returns <code>m/44'/60'/0'/0/2</code>.</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>12</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def bip44_path(coin_type, account, change, index):
    # Purpose is always 44'; coin_type, account are hardened; change, index are not.
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def bip44_path(coin_type, account, change, index):
    return "m/44'/" + str(coin_type) + "'/" + str(account) + "'/" + str(change) + "/" + str(index)
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(bip44_path(60, 0, 0, 2))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>m/44'/60'/0'/0/2
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(bip44_path(0, 3, 1, 14))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>m/44'/0'/3'/1/14
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(bip44_path(60, 0, 0, 0))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>m/44'/60'/0'/0/0
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name><text>S5 Q2 — Hardened or normal?</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S5 Q2 — Hardened indices</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">An index is hardened when it is at least 2<sup>31</sup>.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write <strong><code>is_hardened(index)</code></strong> that returns <code>True</code> if the 32-bit index is hardened (<code>index &gt;= 2**31</code>) and <code>False</code> otherwise.</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>12</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def is_hardened(index):
    # Hardened when index >= 2**31.
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def is_hardened(index):
    return index >= 2**31
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(is_hardened(0))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>False
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(is_hardened(2**31))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>True
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(is_hardened(2147483647))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>False
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(is_hardened(2147483648))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>True
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name><text>S5 Q3 — Display an index</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S5 Q3 — Human-readable indices</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">Wallets show hardened indices counting from zero, followed by a prime mark.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write <strong><code>display_index(index)</code></strong>. For a normal index (below 2<sup>31</sup>) return its number as a string. For a hardened index (2<sup>31</sup> or above) return <code>str(index - 2**31)</code> followed by an apostrophe <code>'</code>. So <code>display_index(0)</code> gives <code>"0"</code>, <code>display_index(2**31)</code> gives <code>"0'"</code>, and <code>display_index(2**31 + 1)</code> gives <code>"1'"</code>.</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>16</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def display_index(index):
    # Normal -> just the number; hardened -> (index - 2**31) followed by an apostrophe.
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def display_index(index):
    if index >= 2**31:
        return str(index - 2**31) + "'"
    return str(index)
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(display_index(0))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>0
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(display_index(2**31))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>0'
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(display_index(2**31 + 1))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>1'
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(display_index(44 + 2**31))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>44'
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>
</quiz>
