<?xml version="1.0" encoding="UTF-8"?>
<quiz><question type="coderunner"><name><text>Q1 — Normalisation et k plus proches voisins</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;K PLUS PROCHES VOISINS&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;La distance euclidienne additionne des carres d'ecarts : une variable exprimee en milliers ecrase toutes les autres. La normalisation est donc obligatoire, et elle doit etre calculee sur les seules donnees d'apprentissage.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Ecrire &lt;strong&gt;&lt;code&gt;normaliser(X_app, X_test)&lt;/code&gt;&lt;/strong&gt; qui renvoie le couple des deux tableaux centres et reduits &lt;em&gt;avec la moyenne et l'ecart-type de &lt;code&gt;X_app&lt;/code&gt; seulement&lt;/em&gt; (un ecart-type nul est remplace par 1), puis &lt;strong&gt;&lt;code&gt;knn(X_app, y_app, X_test, k)&lt;/code&gt;&lt;/strong&gt; qui renvoie le tableau des classes predites par vote majoritaire. En cas d'egalite des votes, choisir la plus petite etiquette. Les tableaux sont des &lt;code&gt;numpy.ndarray&lt;/code&gt;.&lt;/p&gt;&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#014a55;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;X_app = np.array([[0.], [1.], [10.]]) ; y_app = np.array([0, 0, 1])
knn(X_app, y_app, np.array([[9.]]), 1) &amp;rarr; [1]&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def normaliser(X_app, X_test):
    # moyenne et ecart-type calcules sur X_app UNIQUEMENT
    # remplacer un ecart-type nul par 1 pour ne pas diviser par zero
    ...

def knn(X_app, y_app, X_test, k):
    # normaliser, puis pour chaque point de test :
    # distances a tous les points d'apprentissage, tri, vote des k premiers
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

def normaliser(X_app, X_test):
    mu = X_app.mean(axis=0)
    sigma = X_app.std(axis=0)
    sigma = np.where(sigma == 0, 1.0, sigma)
    return (X_app - mu) / sigma, (X_test - mu) / sigma

def knn(X_app, y_app, X_test, k):
    Xa, Xt = normaliser(X_app, X_test)
    etiquettes = np.unique(y_app)
    predictions = []
    for x in Xt:
        d = np.sqrt(((Xa - x) ** 2).sum(axis=1))
        proches = y_app[np.argsort(d, kind='stable')[:k]]
        comptes = [(-(proches == e).sum(), e) for e in etiquettes]
        comptes.sort()
        predictions.append(comptes[0][1])
    return np.array(predictions)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Xa = np.array([[0.], [1.], [10.]]); ya = np.array([0, 0, 1])
print(knn(Xa, ya, np.array([[9.]]), 1))
</text></testcode><stdin><text /></stdin><expected><text>[1]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Xa = np.array([[0., 1000.], [1., 1001.], [10., 1002.]])
a, t = normaliser(Xa, np.array([[5., 1001.]]))
print(np.round(a.mean(axis=0), 6), np.round(a.std(axis=0), 6))
</text></testcode><stdin><text /></stdin><expected><text>[0. 0.] [1. 1.]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Xa = np.array([[0., 0.], [0., 1.], [5., 5.], [5., 6.]]); ya = np.array([0, 0, 1, 1])
print(knn(Xa, ya, np.array([[0.5, 0.5], [4.5, 5.5]]), 3))
</text></testcode><stdin><text /></stdin><expected><text>[0 1]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Xa = np.array([[1.], [1.], [1.]]); ya = np.array([0, 1, 1])
print(knn(Xa, ya, np.array([[1.]]), 3))
</text></testcode><stdin><text /></stdin><expected><text>[1]
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Xa = np.array([[0.], [2.]]); ya = np.array([1, 0])
print(knn(Xa, ya, np.array([[0.5], [1.5]]), 1))
</text></testcode><stdin><text /></stdin><expected><text>[1 0]
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>Q2 — Descente de gradient</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;REGRESSION LINEAIRE PAR GRADIENT&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;La regression lineaire admet une solution fermee, mais la descente de gradient est la seule methode qui se transpose telle quelle a tous les modeles ulterieurs.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Ecrire &lt;strong&gt;&lt;code&gt;descente(X, y, eta, n_iter)&lt;/code&gt;&lt;/strong&gt; qui minimise &lt;code&gt;J(theta) = moyenne((y - X @ theta) ** 2)&lt;/code&gt; par descente de gradient depuis &lt;code&gt;theta = 0&lt;/code&gt;, et renvoie le couple &lt;code&gt;(theta, erreurs)&lt;/code&gt; ou &lt;code&gt;erreurs&lt;/code&gt; est la liste des valeurs de &lt;code&gt;J&lt;/code&gt; &lt;em&gt;avant&lt;/em&gt; chaque mise a jour (donc de longueur &lt;code&gt;n_iter&lt;/code&gt;). Le gradient vaut &lt;code&gt;-2 / n * X.T @ (y - X @ theta)&lt;/code&gt;. On suppose que &lt;code&gt;X&lt;/code&gt; contient deja sa colonne de biais.&lt;/p&gt;&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#014a55;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;X = np.array([[1., 0.], [1., 1.], [1., 2.]]) ; y = np.array([1., 3., 5.])
descente(X, y, 0.1, 500) donne theta proche de [1., 2.]&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def descente(X, y, eta, n_iter):
    # theta initialise a zero
    # a chaque tour : residu, erreur quadratique moyenne, gradient, mise a jour
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

def descente(X, y, eta, n_iter):
    n, d = X.shape
    theta = np.zeros(d)
    erreurs = []
    for _ in range(n_iter):
        residu = y - X @ theta
        erreurs.append(float((residu ** 2).mean()))
        grad = -2.0 / n * (X.T @ residu)
        theta = theta - eta * grad
    return theta, erreurs
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[1., 0.], [1., 1.], [1., 2.]]); y = np.array([1., 3., 5.])
th, err = descente(X, y, 0.1, 500)
print(np.round(th, 3), len(err), round(err[0], 3))
</text></testcode><stdin><text /></stdin><expected><text>[1. 2.] 500 11.667
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[1., 0.], [1., 1.], [1., 2.]]); y = np.array([1., 3., 5.])
th, err = descente(X, y, 0.1, 500)
print(round(err[-1], 8) == 0.0, all(err[i] &gt;= err[i+1] for i in range(len(err)-1)))
</text></testcode><stdin><text /></stdin><expected><text>True True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[1.], [1.], [1.]]); y = np.array([2., 4., 6.])
th, err = descente(X, y, 0.2, 200)
print(np.round(th, 4))
</text></testcode><stdin><text /></stdin><expected><text>[4.]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[1., 1.], [1., 2.], [1., 3.], [1., 4.]])
y = np.array([2.1, 3.9, 6.2, 7.8])
th, _ = descente(X, y, 0.05, 4000)
exact = np.linalg.lstsq(X, y, rcond=None)[0]
print(np.allclose(th, exact, atol=1e-4))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[1., 0.], [1., 1.]]); y = np.array([0., 1.])
th, err = descente(X, y, 0.1, 1)
print(np.round(th, 4), np.round(err, 4))
</text></testcode><stdin><text /></stdin><expected><text>[0.1 0.1] [0.5]
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>Q3 — Entropie et gain d'information</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;CRITERE DE COUPE D'UN ARBRE&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;A chaque noeud d'un arbre de decision, on retient le test qui reduit le plus l'heterogeneite moyenne des groupes obtenus.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Ecrire &lt;strong&gt;&lt;code&gt;entropie(etiquettes)&lt;/code&gt;&lt;/strong&gt; (en bits, avec la convention &lt;code&gt;0 log 0 = 0&lt;/code&gt;, et 0 pour une liste vide), puis &lt;strong&gt;&lt;code&gt;gain(etiquettes, groupes)&lt;/code&gt;&lt;/strong&gt; qui renvoie la difference entre l'entropie de l'ensemble et la moyenne des entropies des groupes ponderee par leurs effectifs. &lt;code&gt;groupes&lt;/code&gt; est une liste de listes d'etiquettes. Arrondir les deux resultats a 4 decimales.&lt;/p&gt;&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#014a55;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;entropie(['oui'] * 9 + ['non'] * 5) &amp;rarr; 0.9403
gain(['oui'] * 9 + ['non'] * 5, [['oui']*2 + ['non']*3, ['oui']*4, ['oui']*3 + ['non']*2]) &amp;rarr; 0.2467&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>from math import log2
from collections import Counter

def entropie(etiquettes):
    # - somme des p * log2(p) sur les classes presentes
    ...

def gain(etiquettes, groupes):
    # entropie du tout, moins la moyenne ponderee des entropies des groupes
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>from math import log2
from collections import Counter

def entropie(etiquettes):
    n = len(etiquettes)
    if n == 0:
        return 0.0
    h = 0.0
    for effectif in Counter(etiquettes).values():
        p = effectif / n
        if p &gt; 0:
            h -= p * log2(p)
    return round(h, 4)

def gain(etiquettes, groupes):
    n = len(etiquettes)
    if n == 0:
        return 0.0
    reste = sum(len(g) / n * entropie(g) for g in groupes)
    return round(entropie(etiquettes) - reste, 4)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(entropie(['oui'] * 9 + ['non'] * 5))
</text></testcode><stdin><text /></stdin><expected><text>0.9403
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(gain(['oui'] * 9 + ['non'] * 5,
           [['oui'] * 2 + ['non'] * 3, ['oui'] * 4, ['oui'] * 3 + ['non'] * 2]))
</text></testcode><stdin><text /></stdin><expected><text>0.2467
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(entropie([]), entropie(['a', 'a', 'a']), entropie(['a', 'b']))
</text></testcode><stdin><text /></stdin><expected><text>0.0 0.0 1.0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(gain(['oui'] * 9 + ['non'] * 5,
           [['oui'] * 6 + ['non'] * 2, ['oui'] * 3 + ['non'] * 3]))
</text></testcode><stdin><text /></stdin><expected><text>0.0481
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(entropie(['a', 'b', 'c', 'd']), gain(['a', 'b'], [['a'], ['b']]))
</text></testcode><stdin><text /></stdin><expected><text>2.0 1.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>Q4 — Matrice de confusion</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;MESURER SANS SE MENTIR&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;Sur des classes desequilibrees, l'exactitude seule est trompeuse : un classifieur qui repond toujours 'negatif' peut afficher 94 % d'exactitude avec un rappel nul.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Ecrire &lt;strong&gt;&lt;code&gt;matrice(y_vrai, y_pred)&lt;/code&gt;&lt;/strong&gt; qui renvoie le quadruplet &lt;code&gt;(vp, fp, fn, vn)&lt;/code&gt; pour la classe positive 1, puis &lt;strong&gt;&lt;code&gt;scores(y_vrai, y_pred)&lt;/code&gt;&lt;/strong&gt; qui renvoie le dictionnaire des cles &lt;code&gt;'exactitude'&lt;/code&gt;, &lt;code&gt;'precision'&lt;/code&gt;, &lt;code&gt;'rappel'&lt;/code&gt; et &lt;code&gt;'f1'&lt;/code&gt;, chaque valeur arrondie a 4 decimales. Une metrique dont le denominateur est nul vaut &lt;code&gt;0.0&lt;/code&gt;. Les entrees sont des listes ou des tableaux d'entiers 0 ou 1.&lt;/p&gt;&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#014a55;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;matrice([1,1,0,0,0], [1,0,1,0,0]) &amp;rarr; (1, 1, 1, 2)
scores([1,1,0,0,0], [1,0,1,0,0])['rappel'] &amp;rarr; 0.5&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>def matrice(y_vrai, y_pred):
    # compter vp, fp, fn, vn pour la classe positive 1
    ...

def scores(y_vrai, y_pred):
    # exactitude = (vp + vn) / total
    # precision  = vp / (vp + fp), rappel = vp / (vp + fn)
    # f1 = moyenne harmonique ; denominateur nul -&gt; 0.0
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>def matrice(y_vrai, y_pred):
    vp = fp = fn = vn = 0
    for a, b in zip(y_vrai, y_pred):
        if a == 1 and b == 1:
            vp += 1
        elif a == 0 and b == 1:
            fp += 1
        elif a == 1 and b == 0:
            fn += 1
        else:
            vn += 1
    return (vp, fp, fn, vn)

def scores(y_vrai, y_pred):
    vp, fp, fn, vn = matrice(y_vrai, y_pred)
    total = vp + fp + fn + vn
    exactitude = (vp + vn) / total if total else 0.0
    precision = vp / (vp + fp) if (vp + fp) else 0.0
    rappel = vp / (vp + fn) if (vp + fn) else 0.0
    f1 = 2 * precision * rappel / (precision + rappel) if (precision + rappel) else 0.0
    return {'exactitude': round(exactitude, 4), 'precision': round(precision, 4),
            'rappel': round(rappel, 4), 'f1': round(f1, 4)}
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(matrice([1,1,0,0,0], [1,0,1,0,0]))
</text></testcode><stdin><text /></stdin><expected><text>(1, 1, 1, 2)
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>s = scores([1,1,0,0,0], [1,0,1,0,0])
print(s['exactitude'], s['precision'], s['rappel'], s['f1'])
</text></testcode><stdin><text /></stdin><expected><text>0.6 0.5 0.5 0.5
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>vrai = [1] * 60 + [0] * 940
pred = [1] * 42 + [0] * 18 + [1] * 27 + [0] * 913
s = scores(vrai, pred)
print(s['exactitude'], s['precision'], s['rappel'], s['f1'])
</text></testcode><stdin><text /></stdin><expected><text>0.955 0.6087 0.7 0.6512
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>vrai = [1] * 60 + [0] * 940
s = scores(vrai, [0] * 1000)
print(s['exactitude'], s['precision'], s['rappel'], s['f1'])
</text></testcode><stdin><text /></stdin><expected><text>0.94 0.0 0.0 0.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(matrice([], []), scores([1], [1]))
</text></testcode><stdin><text /></stdin><expected><text>(0, 0, 0, 0) {'exactitude': 1.0, 'precision': 1.0, 'rappel': 1.0, 'f1': 1.0}
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>Q5 — Validation croisee sans fuite de donnees</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #026573;border-radius:8px;padding:20px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#026573;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;EVALUATION HONNETE&lt;/p&gt;&lt;p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;"&gt;Tout pretraitement qui apprend quelque chose des donnees doit etre ajuste a l'interieur de chaque bloc de validation. Normaliser avant de decouper fait entrer de l'information du bloc de test dans le modele.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,serif;font-size:15px;color:#1e293b;line-height:1.8;"&gt;Ecrire &lt;strong&gt;&lt;code&gt;choisir_k(X, y, ks)&lt;/code&gt;&lt;/strong&gt; qui renvoie le couple &lt;code&gt;(meilleur_k, meilleur_score)&lt;/code&gt;, ou le score est la moyenne des taux de bonnes reponses d'une validation croisee stratifiee a 5 blocs (&lt;code&gt;StratifiedKFold(n_splits=5, shuffle=True, random_state=0)&lt;/code&gt;), arrondie a 4 decimales. Le modele est un enchainement &lt;code&gt;StandardScaler&lt;/code&gt; puis &lt;code&gt;KNeighborsClassifier(n_neighbors=k)&lt;/code&gt; construit avec &lt;code&gt;make_pipeline&lt;/code&gt; : c'est lui qui empeche la fuite. A score egal, garder le plus petit &lt;code&gt;k&lt;/code&gt;.&lt;/p&gt;&lt;div style="background:#e4f1f2;border:1px solid #b7dadd;border-left:5px solid #026573;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#014a55;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;"&gt;sur le jeu iris : choisir_k(X, y, [1, 5, 11]) &amp;rarr; (11, 0.9667)&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import StandardScaler
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import cross_val_score, StratifiedKFold

def choisir_k(X, y, ks):
    # decoupe stratifiee a 5 blocs, melangee, random_state=0
    # pour chaque k : pipeline (normalisation puis voisins), score moyen arrondi
    # garder le meilleur, en conservant le plus petit k a score egal
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>from sklearn.pipeline import make_pipeline
from sklearn.preprocessing import StandardScaler
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import cross_val_score, StratifiedKFold

def choisir_k(X, y, ks):
    decoupe = StratifiedKFold(n_splits=5, shuffle=True, random_state=0)
    meilleur_k, meilleur_score = None, -1.0
    for k in ks:
        modele = make_pipeline(StandardScaler(), KNeighborsClassifier(n_neighbors=k))
        score = round(float(cross_val_score(modele, X, y, cv=decoupe).mean()), 4)
        if score &gt; meilleur_score:
            meilleur_k, meilleur_score = k, score
    return (meilleur_k, meilleur_score)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>from sklearn.datasets import load_iris
X, y = load_iris(return_X_y=True)
print(choisir_k(X, y, [1, 5, 11]))
</text></testcode><stdin><text /></stdin><expected><text>(11, 0.9667)
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>from sklearn.datasets import load_wine
X, y = load_wine(return_X_y=True)
print(choisir_k(X, y, [1, 3, 5, 9, 15]))
</text></testcode><stdin><text /></stdin><expected><text>(15, 0.9665)
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>from sklearn.datasets import load_iris
X, y = load_iris(return_X_y=True)
print(choisir_k(X, y, [7]))
</text></testcode><stdin><text /></stdin><expected><text>(7, 0.96)
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>from sklearn.datasets import load_breast_cancer
X, y = load_breast_cancer(return_X_y=True)
print(choisir_k(X, y, [1, 5, 21]))
</text></testcode><stdin><text /></stdin><expected><text>(5, 0.9649)
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>from sklearn.datasets import load_wine
X, y = load_wine(return_X_y=True)
k, s = choisir_k(X, y, [5])
print(k, s &gt; 0.9)
</text></testcode><stdin><text /></stdin><expected><text>5 True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question></quiz>