<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name><text>S4 Q1 — Base58 encoding</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S4 Q1 — Base58 encoding</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">Base58 uses the alphabet 123456789ABCDEFGHJKLMNPQRSTUVWXYZabcdefghijkmnopqrstuvwxyz. Each leading zero byte becomes a leading '1'.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write <strong><code>base58_encode(data)</code></strong> that encodes a <code>bytes</code> value to a Base58 string. Interpret <code>data</code> as a big-endian integer, and add one leading <code>'1'</code> for every leading zero byte.</p>
<div style="background:#fdf4ff;border-left:5px solid #9333ea;border-radius:8px;padding:12px 18px;margin:12px 0;">
<pre style="color:#3b0764;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;">base58_encode(bytes([0]))    &rarr; '1'
base58_encode(bytes([58]))   &rarr; '21'
base58_encode(bytes([1, 0])) &rarr; '5R'</pre>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[ALPHABET = "123456789ABCDEFGHJKLMNPQRSTUVWXYZabcdefghijkmnopqrstuvwxyz"

def base58_encode(data):
    # Complete: convert bytes -> big-endian int -> base58, then add leading '1's
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[ALPHABET = "123456789ABCDEFGHJKLMNPQRSTUVWXYZabcdefghijkmnopqrstuvwxyz"

def base58_encode(data):
    n = int.from_bytes(data, "big")
    chars = []
    while n > 0:
        n, r = divmod(n, 58)
        chars.append(ALPHABET[r])
    for b in data:
        if b == 0:
            chars.append("1")
        else:
            break
    return "".join(reversed(chars)) if chars else ""
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(base58_encode(bytes([0])))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>1
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(base58_encode(bytes([58])))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>21
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(base58_encode(bytes([1, 0])))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>5R
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(base58_encode(bytes([0, 0, 1])))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>112
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name><text>S4 Q2 — Base58Check checksum</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S4 Q2 — Base58Check checksum</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">The Base58Check checksum is the first 4 bytes of the double SHA-256 of the payload.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Using <code>hashlib</code>, write <strong><code>checksum(payload)</code></strong> that returns the 4-byte checksum (a <code>bytes</code> object) equal to the first four bytes of <code>SHA256(SHA256(payload))</code>.</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[import hashlib

def checksum(payload):
    # Complete: first 4 bytes of double SHA-256
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[import hashlib

def checksum(payload):
    return hashlib.sha256(hashlib.sha256(payload).digest()).digest()[:4]
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(len(checksum(b"abc")))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>4
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>import hashlib
p = bytes([0]) + bytes(20)
exp = hashlib.sha256(hashlib.sha256(p).digest()).digest()[:4]
print(checksum(p) == exp)</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>True
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(checksum(b"a") != checksum(b"b"))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>True
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name><text>S4 Q3 — Compressed public key</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S4 Q3 — Compressed public key</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">A compressed public key is 33 bytes: prefix 0x02 if y is even, 0x03 if y is odd, then the 32-byte big-endian x-coordinate.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write <strong><code>compress_pubkey(x, y)</code></strong> that returns the 33-byte compressed encoding for integer coordinates <code>x</code> and <code>y</code>.</p>
<div style="background:#fdf4ff;border-left:5px solid #9333ea;border-radius:8px;padding:12px 18px;margin:12px 0;">
<pre style="color:#3b0764;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;">compress_pubkey(1, 2).hex()[:2] &rarr; '02'
compress_pubkey(1, 3).hex()[:2] &rarr; '03'</pre>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def compress_pubkey(x, y):
    prefix = 0x02 if y % 2 == 0 else 0x03
    # Complete: prefix byte followed by 32-byte big-endian x
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def compress_pubkey(x, y):
    prefix = 0x02 if y % 2 == 0 else 0x03
    return bytes([prefix]) + x.to_bytes(32, "big")
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(len(compress_pubkey(1, 2)))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>33
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(compress_pubkey(1, 2).hex()[:2])</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>02
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(compress_pubkey(1, 3).hex()[:2])</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>03
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(compress_pubkey(1, 2).hex()[-2:])</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>01
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
    </testcases>
  </question>
</quiz>
