<?xml version="1.0" encoding="UTF-8"?>
<quiz><question type="coderunner"><name><text>C3.Q1 — A state-vector simulator by tensor contraction</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C3.Q1 — Applying gates without building 2^n x 2^n matrices&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;Building the full \(2^{n}\times2^{n}\) matrix of a gate wastes exponential memory. Reshaping the state to a rank-\(n\) tensor and contracting only the relevant axes costs \(O(2^{n})\) per gate.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;apply_1q(state, U, q, n)&lt;/code&gt; and &lt;code&gt;apply_2q(state, U, q0, q1, n)&lt;/code&gt; as described in §5, using &lt;code&gt;numpy.tensordot&lt;/code&gt; and &lt;code&gt;numpy.moveaxis&lt;/code&gt;. Convention: qubit 0 is the &lt;strong&gt;most significant&lt;/strong&gt; factor. Both return a flat length-\(2^{n}\) array.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;H on qubit 0 of |00&amp;gt;, then CNOT(0,1)  -&amp;gt;  (|00&amp;gt; + |11&amp;gt;)/sqrt(2)&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def apply_1q(state, U, q, n):
    # reshape -&gt; tensordot -&gt; moveaxis -&gt; reshape
    ...

def apply_2q(state, U, q0, q1, n):
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

def apply_1q(state, U, q, n):
    psi = np.asarray(state, dtype=complex).reshape([2] * n)
    psi = np.tensordot(np.asarray(U, dtype=complex), psi, axes=([1], [q]))
    psi = np.moveaxis(psi, 0, q)
    return psi.reshape(2 ** n)

def apply_2q(state, U, q0, q1, n):
    psi = np.asarray(state, dtype=complex).reshape([2] * n)
    U4 = np.asarray(U, dtype=complex).reshape(2, 2, 2, 2)
    psi = np.tensordot(U4, psi, axes=([2, 3], [q0, q1]))
    psi = np.moveaxis(psi, [0, 1], [q0, q1])
    return psi.reshape(2 ** n)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
H = np.array([[1,1],[1,-1]], dtype=complex)/np.sqrt(2)
CNOT = np.array([[1,0,0,0],[0,1,0,0],[0,0,0,1],[0,0,1,0]], dtype=complex)
psi = np.zeros(4, dtype=complex); psi[0] = 1
psi = apply_1q(psi, H, 0, 2)
psi = apply_2q(psi, CNOT, 0, 1, 2)
print(np.round(psi.real, 6).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[0.707107, 0.0, 0.0, 0.707107]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[0,1],[1,0]], dtype=complex)
psi = np.zeros(8, dtype=complex); psi[0] = 1
psi = apply_1q(psi, X, 2, 3)
print(int(np.argmax(np.abs(psi))))
</text></testcode><stdin><text /></stdin><expected><text>1
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
CNOT = np.array([[1,0,0,0],[0,1,0,0],[0,0,0,1],[0,0,1,0]], dtype=complex)
psi = np.zeros(4, dtype=complex); psi[2] = 1     # |10&gt;
psi = apply_2q(psi, CNOT, 0, 1, 2)
print(int(np.argmax(np.abs(psi))))               # |11&gt; = index 3
</text></testcode><stdin><text /></stdin><expected><text>3
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
CNOT = np.array([[1,0,0,0],[0,1,0,0],[0,0,0,1],[0,0,1,0]], dtype=complex)
psi = np.zeros(4, dtype=complex); psi[2] = 1     # |10&gt;
psi = apply_2q(psi, CNOT, 1, 0, 2)               # control = qubit 1
print(int(np.argmax(np.abs(psi))))               # unchanged: |10&gt; = index 2
</text></testcode><stdin><text /></stdin><expected><text>2
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
rng = np.random.default_rng(3)
H = np.array([[1,1],[1,-1]], dtype=complex)/np.sqrt(2)
psi = rng.normal(size=16) + 1j*rng.normal(size=16)
psi = psi/np.linalg.norm(psi)
out = apply_1q(apply_1q(psi, H, 2, 4), H, 2, 4)
print(np.allclose(out, psi))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C3.Q2 — Controlled gates and the CNOT/CZ relation</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C3.Q2 — Building two-qubit gates&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;\(C\text{-}U=|0\rangle\langle0|\otimes I+|1\rangle\langle1|\otimes U\) in the ordering where the control is the first factor.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;controlled(U)&lt;/code&gt; returning the \(4\times4\) controlled-\(U\); and &lt;code&gt;cnot_from_cz()&lt;/code&gt; returning the CNOT matrix built &lt;em&gt;only&lt;/em&gt; as \((I\otimes H)\,\mathrm{CZ}\,(I\otimes H)\), where CZ is constructed with &lt;code&gt;controlled&lt;/code&gt;.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;controlled(X)  ==  CNOT
cnot_from_cz() ==  CNOT&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def controlled(U):
    # |0&gt;&lt;0| (x) I  +  |1&gt;&lt;1| (x) U
    ...

def cnot_from_cz():
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

def controlled(U):
    U = np.asarray(U, dtype=complex)
    P0 = np.array([[1, 0], [0, 0]], dtype=complex)
    P1 = np.array([[0, 0], [0, 1]], dtype=complex)
    return np.kron(P0, np.eye(2, dtype=complex)) + np.kron(P1, U)

def cnot_from_cz():
    H = np.array([[1, 1], [1, -1]], dtype=complex) / np.sqrt(2)
    Z = np.array([[1, 0], [0, -1]], dtype=complex)
    CZ = controlled(Z)
    IH = np.kron(np.eye(2, dtype=complex), H)
    return IH @ CZ @ IH
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[0,1],[1,0]], dtype=complex)
CNOT = np.array([[1,0,0,0],[0,1,0,0],[0,0,0,1],[0,0,1,0]], dtype=complex)
print(np.allclose(controlled(X), CNOT))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
CNOT = np.array([[1,0,0,0],[0,1,0,0],[0,0,0,1],[0,0,1,0]], dtype=complex)
print(np.allclose(cnot_from_cz(), CNOT))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Z = np.array([[1,0],[0,-1]], dtype=complex)
print(np.round(np.diag(controlled(Z)).real, 6).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[1.0, 1.0, 1.0, -1.0]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
T = np.diag([1, np.exp(1j*np.pi/4)]).astype(complex)
CT = controlled(T)
print(np.allclose(CT.conj().T @ CT, np.eye(4)), np.allclose(np.linalg.matrix_power(CT, 8), np.eye(4)))
</text></testcode><stdin><text /></stdin><expected><text>True True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Z = np.array([[1,0],[0,-1]], dtype=complex)
CZ = controlled(Z)
S = np.array([[0,0,1,0],[0,1,0,0],[1,0,0,0],[0,0,0,1]], dtype=complex)  # swap basis 00&lt;-&gt;10
print(np.allclose(CZ, S @ CZ @ S))   # CZ is symmetric under qubit exchange
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C3.Q3 — ZYZ (Euler) decomposition of a single-qubit unitary</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C3.Q3 — Compiling an arbitrary gate&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;Every \(U\in U(2)\) factorises as \(e^{i\alpha}R_z(\beta)R_y(\gamma)R_z(\delta)\). Extracting the angles is the first step of every quantum compiler.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;zyz(U)&lt;/code&gt; returning the tuple &lt;code&gt;(alpha, beta, gamma, delta)&lt;/code&gt; of floats, and &lt;code&gt;from_zyz(alpha, beta, gamma, delta)&lt;/code&gt; rebuilding the matrix, with \(R_z(\lambda)=\mathrm{diag}(e^{-i\lambda/2},e^{i\lambda/2})\) and \(R_y(\gamma)=\begin{pmatrix}\cos\frac\gamma2&amp;-\sin\frac\gamma2\\\sin\frac\gamma2&amp;\cos\frac\gamma2\end{pmatrix}\). Grading only checks that &lt;code&gt;from_zyz(*zyz(U))&lt;/code&gt; reproduces &lt;code&gt;U&lt;/code&gt;.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;from_zyz(*zyz(H))  ==  H     (up to numerical tolerance)&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def rz(l):
    ...

def ry(g):
    ...

def from_zyz(alpha, beta, gamma, delta):
    ...

def zyz(U):
    # hint: remove the global phase via det(U), then read the angles off the SU(2) matrix
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

def rz(l):
    return np.diag([np.exp(-1j * l / 2), np.exp(1j * l / 2)])

def ry(g):
    c, s = np.cos(g / 2), np.sin(g / 2)
    return np.array([[c, -s], [s, c]], dtype=complex)

def from_zyz(alpha, beta, gamma, delta):
    return np.exp(1j * alpha) * (rz(beta) @ ry(gamma) @ rz(delta))

def zyz(U):
    U = np.asarray(U, dtype=complex)
    det = np.linalg.det(U)
    alpha = float(np.angle(det) / 2)
    V = U * np.exp(-1j * alpha)                 # V in SU(2)
    gamma = float(2 * np.arctan2(abs(V[1, 0]), abs(V[0, 0])))
    # V = [[e^{-i(b+d)/2} cos(g/2), -e^{-i(b-d)/2} sin(g/2)],
    #      [ e^{i(b-d)/2} sin(g/2),  e^{i(b+d)/2} cos(g/2)]]
    half_sum = np.angle(V[1, 1]) if abs(V[1, 1]) &gt; 1e-12 else 0.0   # (b+d)/2
    half_dif = np.angle(V[1, 0]) if abs(V[1, 0]) &gt; 1e-12 else 0.0   # (b-d)/2
    beta = float(half_sum + half_dif)
    delta = float(half_sum - half_dif)
    return (alpha, beta, gamma, delta)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
H = np.array([[1,1],[1,-1]], dtype=complex)/np.sqrt(2)
print(np.allclose(from_zyz(*zyz(H)), H, atol=1e-8))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
T = np.diag([1, np.exp(1j*np.pi/4)]).astype(complex)
print(np.allclose(from_zyz(*zyz(T)), T, atol=1e-8))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[0,1],[1,0]], dtype=complex)
print(np.allclose(from_zyz(*zyz(X)), X, atol=1e-8))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
rng = np.random.default_rng(2024)
ok = True
for _ in range(20):
    A = rng.normal(size=(2,2)) + 1j*rng.normal(size=(2,2))
    Q, R = np.linalg.qr(A)
    Q = Q @ np.diag(np.exp(1j*np.angle(np.diag(R))))
    ok = ok and np.allclose(from_zyz(*zyz(Q)), Q, atol=1e-7)
print(ok)
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
I = np.eye(2, dtype=complex)
print(np.allclose(from_zyz(*zyz(I)), I, atol=1e-8))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C3.Q4 — Running a circuit: Bell, GHZ and SWAP from CNOTs</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C3.Q4 — A minimal circuit runner&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;A circuit is a list of instructions. This exercise assembles the primitives of Q1–Q2 into a runner and verifies three standard constructions.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;run(n, circuit)&lt;/code&gt; where &lt;code&gt;circuit&lt;/code&gt; is a list of tuples &lt;code&gt;(gate_name, qubits)&lt;/code&gt; with &lt;code&gt;gate_name&lt;/code&gt; in &lt;code&gt;{'H','X','Z','S','T','CNOT','CZ','SWAP'}&lt;/code&gt; and &lt;code&gt;qubits&lt;/code&gt; a tuple of indices (control first for CNOT/CZ). Start from \(|0\rangle^{\otimes n}\) and return the final state vector. Implement SWAP as three CNOTs.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;run(2, [('H',(0,)), ('CNOT',(0,1))])  -&amp;gt;  Bell state
run(3, [('H',(0,)), ('CNOT',(0,1)), ('CNOT',(1,2))])  -&amp;gt;  GHZ&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def run(n, circuit):
    # start from |0...0&gt;, apply each instruction in order
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

H = np.array([[1, 1], [1, -1]], dtype=complex) / np.sqrt(2)
X = np.array([[0, 1], [1, 0]], dtype=complex)
Z = np.array([[1, 0], [0, -1]], dtype=complex)
S = np.diag([1, 1j]).astype(complex)
T = np.diag([1, np.exp(1j * np.pi / 4)]).astype(complex)
CNOT = np.array([[1,0,0,0],[0,1,0,0],[0,0,0,1],[0,0,1,0]], dtype=complex)
CZ = np.diag([1, 1, 1, -1]).astype(complex)
ONE_Q = {'H': H, 'X': X, 'Z': Z, 'S': S, 'T': T}
TWO_Q = {'CNOT': CNOT, 'CZ': CZ}

def _apply1(state, U, q, n):
    psi = state.reshape([2] * n)
    psi = np.tensordot(U, psi, axes=([1], [q]))
    return np.moveaxis(psi, 0, q).reshape(2 ** n)

def _apply2(state, U, q0, q1, n):
    psi = state.reshape([2] * n)
    psi = np.tensordot(U.reshape(2, 2, 2, 2), psi, axes=([2, 3], [q0, q1]))
    return np.moveaxis(psi, [0, 1], [q0, q1]).reshape(2 ** n)

def run(n, circuit):
    psi = np.zeros(2 ** n, dtype=complex)
    psi[0] = 1
    for name, qs in circuit:
        if name in ONE_Q:
            psi = _apply1(psi, ONE_Q[name], qs[0], n)
        elif name in TWO_Q:
            psi = _apply2(psi, TWO_Q[name], qs[0], qs[1], n)
        elif name == 'SWAP':
            a, b = qs
            psi = _apply2(psi, CNOT, a, b, n)
            psi = _apply2(psi, CNOT, b, a, n)
            psi = _apply2(psi, CNOT, a, b, n)
        else:
            raise ValueError('unknown gate ' + name)
    return psi
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
psi = run(2, [('H',(0,)), ('CNOT',(0,1))])
print(np.round(psi.real, 6).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[0.707107, 0.0, 0.0, 0.707107]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
psi = run(3, [('H',(0,)), ('CNOT',(0,1)), ('CNOT',(1,2))])
print(np.round(np.abs(psi)**2, 6).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[0.5, 0.0, 0.0, 0.0, 0.0, 0.0, 0.0, 0.5]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
psi = run(2, [('X',(0,)), ('SWAP',(0,1))])
print(int(np.argmax(np.abs(psi))))
</text></testcode><stdin><text /></stdin><expected><text>1
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
psi = run(1, [('H',(0,)), ('T',(0,)), ('T',(0,)), ('H',(0,))])
ref = np.array([1+1j, 1-1j])/2
print(np.allclose(np.abs(psi), np.abs(ref)))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
psi = run(4, [('H',(0,)), ('H',(1,)), ('H',(2,)), ('H',(3,))])
print(np.allclose(np.abs(psi)**2, np.full(16, 1/16)))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question></quiz>