<?xml version="1.0" encoding="UTF-8"?>
<quiz><question type="coderunner"><name><text>C2.Q1 — Bloch vector of a pure qubit state</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C2.Q1 — From amplitudes to the Bloch sphere&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;The Bloch vector \(\vec r=(\langle X\rangle,\langle Y\rangle,\langle Z\rangle)\) is the coordinate system in which single-qubit gates become rotations.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;bloch_vector(psi)&lt;/code&gt; returning a tuple of three &lt;code&gt;float&lt;/code&gt;s (each rounded to 6 decimals) for a normalised 2-component state vector, and &lt;code&gt;state_from_angles(theta, phi)&lt;/code&gt; returning the NumPy array \((\cos\tfrac\theta2,\ e^{i\varphi}\sin\tfrac\theta2)\).&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;bloch_vector([1, 0])             -&amp;gt; (0.0, 0.0, 1.0)
bloch_vector([1, 1]/sqrt(2))     -&amp;gt; (1.0, 0.0, 0.0)
bloch_vector([1, 1j]/sqrt(2))    -&amp;gt; (0.0, 1.0, 0.0)&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def bloch_vector(psi):
    # &lt;X&gt;, &lt;Y&gt;, &lt;Z&gt; rounded to 6 decimals
    ...

def state_from_angles(theta, phi):
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

X = np.array([[0, 1], [1, 0]], dtype=complex)
Y = np.array([[0, -1j], [1j, 0]], dtype=complex)
Z = np.array([[1, 0], [0, -1]], dtype=complex)

def bloch_vector(psi):
    psi = np.asarray(psi, dtype=complex).reshape(2)
    psi = psi / np.linalg.norm(psi)
    return tuple(round(float(np.real(psi.conj() @ (S @ psi))), 6) for S in (X, Y, Z))

def state_from_angles(theta, phi):
    return np.array([np.cos(theta / 2), np.exp(1j * phi) * np.sin(theta / 2)], dtype=complex)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(bloch_vector([1, 0]))
</text></testcode><stdin><text /></stdin><expected><text>(0.0, 0.0, 1.0)
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(bloch_vector(np.array([1, 1])/np.sqrt(2)))
</text></testcode><stdin><text /></stdin><expected><text>(1.0, 0.0, 0.0)
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(bloch_vector(np.array([1, 1j])/np.sqrt(2)))
</text></testcode><stdin><text /></stdin><expected><text>(0.0, 1.0, 0.0)
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
th, ph = 0.9, 2.1
r = bloch_vector(state_from_angles(th, ph))
exp = (np.sin(th)*np.cos(ph), np.sin(th)*np.sin(ph), np.cos(th))
print(np.allclose(r, exp, atol=1e-6))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
rng = np.random.default_rng(11)
v = rng.normal(size=2) + 1j*rng.normal(size=2)
r = np.array(bloch_vector(v))
print(round(float(np.linalg.norm(r)), 5))
</text></testcode><stdin><text /></stdin><expected><text>1.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C2.Q2 — Rotations act as SO(3) on the Bloch ball</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C2.Q2 — Verifying the SU(2) → SO(3) homomorphism&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;\(R_{\hat n}(\theta)=\cos\tfrac\theta2 I-i\sin\tfrac\theta2(\hat n\cdot\vec\sigma)\) rotates the Bloch vector by \(\theta\) about \(\hat n\). This exercise checks that claim numerically.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;rot(n, theta)&lt;/code&gt; returning the \(2\times2\) unitary for a unit axis &lt;code&gt;n = (nx, ny, nz)&lt;/code&gt;, and &lt;code&gt;rotate_bloch(n, theta, r)&lt;/code&gt; returning the classical Rodrigues rotation of the 3-vector &lt;code&gt;r&lt;/code&gt;. Both must agree.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;rot((0,0,1), pi)  ==  -1j * Z
rotate_bloch((0,0,1), pi, (1,0,0))  -&amp;gt;  (-1, 0, 0)&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def rot(n, theta):
    # cos(theta/2) I - i sin(theta/2) (n . sigma)
    ...

def rotate_bloch(n, theta, r):
    # Rodrigues formula
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

X = np.array([[0, 1], [1, 0]], dtype=complex)
Y = np.array([[0, -1j], [1j, 0]], dtype=complex)
Z = np.array([[1, 0], [0, -1]], dtype=complex)
I2 = np.eye(2, dtype=complex)

def rot(n, theta):
    n = np.asarray(n, dtype=float)
    n = n / np.linalg.norm(n)
    ns = n[0] * X + n[1] * Y + n[2] * Z
    return np.cos(theta / 2) * I2 - 1j * np.sin(theta / 2) * ns

def rotate_bloch(n, theta, r):
    n = np.asarray(n, dtype=float); n = n / np.linalg.norm(n)
    r = np.asarray(r, dtype=float)
    return (r * np.cos(theta) + np.cross(n, r) * np.sin(theta)
            + n * np.dot(n, r) * (1 - np.cos(theta)))
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Z = np.array([[1,0],[0,-1]], dtype=complex)
print(np.allclose(rot((0,0,1), np.pi), -1j*Z))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(np.round(rotate_bloch((0,0,1), np.pi, (1,0,0)), 6).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[-1.0, 0.0, 0.0]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
U = rot((1,1,0), 0.83)
print(np.allclose(U.conj().T @ U, np.eye(2)))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[0,1],[1,0]], dtype=complex)
Y = np.array([[0,-1j],[1j,0]], dtype=complex)
Z = np.array([[1,0],[0,-1]], dtype=complex)
def bv(psi):
    psi = psi/np.linalg.norm(psi)
    return np.array([np.real(psi.conj() @ (S @ psi)) for S in (X,Y,Z)])
psi = np.array([np.cos(0.4), np.exp(1j*1.1)*np.sin(0.4)], dtype=complex)
n, th = (0,1,0), 0.77
lhs = bv(rot(n, th) @ psi)
rhs = rotate_bloch(n, th, bv(psi))
print(np.allclose(lhs, rhs, atol=1e-8))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(np.allclose(rot((1,0,0), 2*np.pi), -np.eye(2)))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C2.Q3 — Projective measurement in an arbitrary basis</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C2.Q3 — Born rule and state collapse&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;Given an orthonormal basis \(\{|b_k\rangle\}\), outcome \(k\) occurs with probability \(|\langle b_k|\psi\rangle|^{2}\) and the state collapses to \(|b_k\rangle\).&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;born_probs(psi, basis)&lt;/code&gt; where &lt;code&gt;basis&lt;/code&gt; is a matrix whose &lt;em&gt;columns&lt;/em&gt; are the basis vectors, returning a NumPy array of probabilities (rounded to 8 decimals, summing to 1); and &lt;code&gt;post_measurement(psi, basis, k)&lt;/code&gt; returning the normalised collapsed state.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;psi = |0&amp;gt;, basis = columns [|+&amp;gt;, |-&amp;gt;]  -&amp;gt;  [0.5, 0.5]&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def born_probs(psi, basis):
    # amplitudes = basis^dagger @ psi
    ...

def post_measurement(psi, basis, k):
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

def born_probs(psi, basis):
    psi = np.asarray(psi, dtype=complex).ravel()
    psi = psi / np.linalg.norm(psi)
    B = np.asarray(basis, dtype=complex)
    amps = B.conj().T @ psi
    return np.round(np.abs(amps) ** 2, 8)

def post_measurement(psi, basis, k):
    B = np.asarray(basis, dtype=complex)
    v = B[:, k]
    return v / np.linalg.norm(v)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
H = np.array([[1,1],[1,-1]], dtype=complex)/np.sqrt(2)
print(born_probs([1,0], H).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[0.5, 0.5]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(born_probs([1,0], np.eye(2, dtype=complex)).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[1.0, 0.0]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
H = np.array([[1,1],[1,-1]], dtype=complex)/np.sqrt(2)
v = post_measurement([1,0], H, 1)
print(np.allclose(v, np.array([1,-1])/np.sqrt(2)))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
rng = np.random.default_rng(5)
A = rng.normal(size=(4,4)) + 1j*rng.normal(size=(4,4))
B, _ = np.linalg.qr(A)
psi = rng.normal(size=4) + 1j*rng.normal(size=4)
print(round(float(born_probs(psi, B).sum()), 6))
</text></testcode><stdin><text /></stdin><expected><text>1.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Y = np.array([[1,1],[1j,-1j]], dtype=complex)/np.sqrt(2)
print(np.round(born_probs(np.array([1,1j])/np.sqrt(2), Y), 6).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[1.0, 0.0]
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C2.Q4 — Sampling, shot noise and the Helstrom bound</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C2.Q4 — What a quantum computer actually outputs&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;A device returns samples, not amplitudes. Estimating \(\langle Z\rangle\) from \(N\) shots incurs an error \(\Theta(1/\sqrt N)\); distinguishing non-orthogonal states is capped by the Helstrom bound.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;sample_counts(psi, shots, seed)&lt;/code&gt; returning a dict mapping bit-strings (MSB = qubit 0, zero-padded to \(n\) characters) to counts, sorted by key, using &lt;code&gt;numpy.random.default_rng(seed).choice&lt;/code&gt;; and &lt;code&gt;helstrom(psi, phi)&lt;/code&gt; returning the optimal success probability \(\tfrac12(1+\sqrt{1-|\langle\psi|\phi\rangle|^{2}})\), rounded to 6 decimals.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;helstrom(|0&amp;gt;, |0&amp;gt;)  -&amp;gt;  0.5      (indistinguishable)
helstrom(|0&amp;gt;, |1&amp;gt;)  -&amp;gt;  1.0      (orthogonal)
helstrom(|0&amp;gt;, |+&amp;gt;)  -&amp;gt;  0.853553&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def sample_counts(psi, shots, seed):
    ...

def helstrom(psi, phi):
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

def sample_counts(psi, shots, seed):
    psi = np.asarray(psi, dtype=complex).ravel()
    probs = np.abs(psi) ** 2
    probs = probs / probs.sum()
    n = int(round(np.log2(len(probs))))
    rng = np.random.default_rng(seed)
    idx = rng.choice(len(probs), size=shots, p=probs)
    counts = {}
    for i in idx:
        key = format(int(i), '0{}b'.format(n))
        counts[key] = counts.get(key, 0) + 1
    return dict(sorted(counts.items()))

def helstrom(psi, phi):
    a = np.asarray(psi, dtype=complex).ravel(); a = a / np.linalg.norm(a)
    b = np.asarray(phi, dtype=complex).ravel(); b = b / np.linalg.norm(b)
    ov = abs(np.vdot(a, b)) ** 2
    return round(0.5 * (1 + np.sqrt(max(0.0, 1 - ov))), 6)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(helstrom([1,0], [0,1]), helstrom([1,0], [1,0]))
</text></testcode><stdin><text /></stdin><expected><text>1.0 0.5
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(helstrom([1,0], np.array([1,1])/np.sqrt(2)))
</text></testcode><stdin><text /></stdin><expected><text>0.853553
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
c = sample_counts(np.array([1,1])/np.sqrt(2), 1000, 0)
print(sorted(c.keys()), sum(c.values()))
</text></testcode><stdin><text /></stdin><expected><text>['0', '1'] 1000
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
psi = np.zeros(8, dtype=complex); psi[0]=1/np.sqrt(2); psi[7]=1/np.sqrt(2)
c = sample_counts(psi, 500, 42)
print(sorted(c.keys()), sum(c.values()))
</text></testcode><stdin><text /></stdin><expected><text>['000', '111'] 500
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
psi = np.array([1,0], dtype=complex)
print(sample_counts(psi, 7, 1))
</text></testcode><stdin><text /></stdin><expected><text>{'0': 7}
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question></quiz>