<?xml version="1.0" encoding="UTF-8"?>
<!--
  SCHÉMA EXACT d'UNE question CodeRunner Python (structure validée par export Moodle réel).
  Un exercises.xml = <quiz> contenant plusieurs <question type="coderunner">.
  Règles d'or :
   - <answer> et <answerpreload> contiennent le code DIRECTEMENT (via CDATA), PAS de <text>.
   - <questiontext> utilise <text><![CDATA[ ... HTML ... ]]></text>.
   - <display> est un ÉLÉMENT ENFANT avec <text>, JAMAIS un attribut de <testcase>.
   - Terminer chaque <expected><text> par un saut de ligne réel.
   - Valeurs fixes : coderunnertype=python3, penaltyregime=0, extractcodefromjson=1,
     templateparamslang=None, giveupallowed=0, hoisttemplateparams=1, defaultgrade=2.0000000.
   - Champs vides requis : idnumber, globalextra, useace, resultcolumns, template,
     iscombinatortemplate, allowmultiplestdins, testsplitterre, language, sandbox, grader,
     uiplugin, uiparameters, prototypeextra.
-->
<quiz>
  <question type="coderunner">
    <name><text>S1 Q1 — Titre court</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S1 Q1 — Titre</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">Contexte de l'exercice.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Écrire <strong><code>ma_fonction(x)</code></strong> qui ...</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def ma_fonction(x):
    # Compléter
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def ma_fonction(x):
    return x
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(ma_fonction(3))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>3
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(ma_fonction(0))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>0
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>
</quiz>
