<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name><text>S1 Q1 — Satoshi conversion</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S1 Q1 — Satoshi conversion</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">One bitcoin equals 100,000,000 satoshis. Protocol accounting is done in integer satoshis.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write a function <strong><code>to_satoshis(btc)</code></strong> that converts an amount in BTC (a float or int) to an <strong>integer</strong> number of satoshis, rounded to the nearest satoshi.</p>
<div style="background:#fdf4ff;border-left:5px solid #9333ea;border-radius:8px;padding:12px 18px;margin:12px 0;">
<pre style="color:#3b0764;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;">to_satoshis(1)          &rarr; 100000000
to_satoshis(0.5)        &rarr; 50000000
to_satoshis(0.00000001) &rarr; 1</pre>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def to_satoshis(btc):
    # Complete: return the integer number of satoshis
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def to_satoshis(btc):
    return round(btc * 100_000_000)
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(to_satoshis(1))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>100000000
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(to_satoshis(0.5))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>50000000
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(to_satoshis(0.00000001))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>1
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(to_satoshis(21_000_000))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>2100000000000000
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name><text>S1 Q2 — Block subsidy and halving</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S1 Q2 — Block subsidy</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">The block subsidy starts at 50 BTC (5,000,000,000 satoshis) and halves every 210,000 blocks. After 64 halvings it reaches 0.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write a function <strong><code>block_subsidy(height)</code></strong> that returns the subsidy in <strong>satoshis</strong> for a block at the given height, using integer right-shift arithmetic (exactly as Bitcoin Core does).</p>
<div style="background:#fdf4ff;border-left:5px solid #9333ea;border-radius:8px;padding:12px 18px;margin:12px 0;">
<pre style="color:#3b0764;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;">block_subsidy(0)      &rarr; 5000000000
block_subsidy(210000) &rarr; 2500000000
block_subsidy(630000) &rarr; 625000000</pre>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def block_subsidy(height):
    halvings = height // 210000
    # Complete: return 0 after 64 halvings, else shift the 50 BTC subsidy
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def block_subsidy(height):
    halvings = height // 210000
    if halvings >= 64:
        return 0
    return 5_000_000_000 >> halvings
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(block_subsidy(0))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>5000000000
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(block_subsidy(210000))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>2500000000
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(block_subsidy(630000))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>625000000
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(block_subsidy(13440000))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>0
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name><text>S1 Q3 — Total money supply</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S1 Q3 — Total supply</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">Summing 210,000 blocks at each subsidy level (using the same integer shift arithmetic) gives the exact number of satoshis that will ever exist.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write a function <strong><code>total_supply()</code></strong> that returns the total number of satoshis ever created, by looping over every halving until the subsidy reaches 0.</p>
<div style="background:#fdf4ff;border-left:5px solid #9333ea;border-radius:8px;padding:12px 18px;margin:12px 0;">
<pre style="color:#3b0764;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;">total_supply() &rarr; 2099999997690000   # just under 21,000,000 BTC</pre>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def total_supply():
    total = 0
    subsidy = 5_000_000_000
    # Complete: accumulate 210000 * subsidy while subsidy > 0, halving each round
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def total_supply():
    total = 0
    subsidy = 5_000_000_000
    while subsidy > 0:
        total += 210000 * subsidy
        subsidy //= 2
    return total
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(total_supply())</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>2099999997690000
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(total_supply() &lt;= 21_000_000 * 100_000_000)</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>True
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(total_supply() // 100_000_000)</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>20999999
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>
</quiz>
