<?xml version="1.0" encoding="UTF-8"?>
<quiz><question type="coderunner"><name><text>C10.Q1 — Pauli-sum Hamiltonians and expectation values</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C10.Q1 — The input format of every variational algorithm&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;A Hamiltonian is given as a list of (Pauli string, coefficient) pairs. Its expectation value on a state is the weighted sum of the individual Pauli expectations — exactly what a device measures.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;pauli_string(s)&lt;/code&gt; building the matrix of a string over &lt;code&gt;{'I','X','Y','Z'}&lt;/code&gt;; &lt;code&gt;hamiltonian(terms)&lt;/code&gt; summing the weighted terms; &lt;code&gt;expectation(psi, H)&lt;/code&gt; returning \(\langle\psi|H|\psi\rangle\) as a real float rounded to 6 decimals; and &lt;code&gt;ground_energy(H)&lt;/code&gt; returning the smallest eigenvalue, rounded to 6 decimals.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;hamiltonian([('ZZ',1.0)]) -&amp;gt; diag(1,-1,-1,1)
ground_energy(hamiltonian([('ZZ',1.0),('XI',0.5),('IX',0.5)])) -&amp;gt; -1.414214&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def pauli_string(s):
    ...

def hamiltonian(terms):
    ...

def expectation(psi, H):
    ...

def ground_energy(H):
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

I2 = np.eye(2, dtype=complex)
X = np.array([[0, 1], [1, 0]], dtype=complex)
Y = np.array([[0, -1j], [1j, 0]], dtype=complex)
Z = np.array([[1, 0], [0, -1]], dtype=complex)
PAULI = {'I': I2, 'X': X, 'Y': Y, 'Z': Z}

def pauli_string(s):
    out = np.array([[1.0 + 0j]])
    for ch in s:
        out = np.kron(out, PAULI[ch])
    return out

def hamiltonian(terms):
    n = len(terms[0][0])
    H = np.zeros((2 ** n, 2 ** n), dtype=complex)
    for s, c in terms:
        H = H + c * pauli_string(s)
    return H

def expectation(psi, H):
    psi = np.asarray(psi, dtype=complex).ravel()
    psi = psi / np.linalg.norm(psi)
    return round(float(np.real(np.vdot(psi, np.asarray(H, dtype=complex) @ psi))), 6)

def ground_energy(H):
    H = np.asarray(H, dtype=complex)
    return round(float(np.linalg.eigvalsh((H + H.conj().T) / 2)[0]), 6)
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(np.round(np.diag(hamiltonian([('ZZ',1.0)])).real, 6).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[1.0, -1.0, -1.0, 1.0]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
H = hamiltonian([('ZZ',1.0),('XI',0.5),('IX',0.5)])
print(ground_energy(H))
</text></testcode><stdin><text /></stdin><expected><text>-1.414214
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
H = hamiltonian([('ZZ',1.0)])
psi = np.array([0,1,0,0], dtype=complex)
print(expectation(psi, H))
</text></testcode><stdin><text /></stdin><expected><text>-1.0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(np.allclose(pauli_string('XY'), np.kron(np.array([[0,1],[1,0]]), np.array([[0,-1j],[1j,0]]))))
</text></testcode><stdin><text /></stdin><expected><text>True
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
H = hamiltonian([('XX',1.0),('YY',1.0),('ZZ',1.0)])
print(ground_energy(H))
</text></testcode><stdin><text /></stdin><expected><text>-3.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C10.Q2 — VQE with the parameter-shift rule</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C10.Q2 — A complete variational loop&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;For \(U(\theta)=e^{-i\theta P/2}\) with \(P^{2}=I\), \(\partial_\theta\langle A\rangle=\tfrac12[\langle A\rangle_{\theta+\pi/2}-\langle A\rangle_{\theta-\pi/2}]\) exactly.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Take the single-qubit ansatz \(|\psi(\theta)\rangle=R_y(\theta)|0\rangle\) with \(R_y(\theta)=\begin{pmatrix}\cos\frac\theta2&amp;-\sin\frac\theta2\\\sin\frac\theta2&amp;\cos\frac\theta2\end{pmatrix}\). Write &lt;code&gt;energy(theta, H)&lt;/code&gt;; &lt;code&gt;grad_parameter_shift(theta, H)&lt;/code&gt; using the exact rule; and &lt;code&gt;vqe(H, theta0, lr=0.3, steps=300)&lt;/code&gt; returning &lt;code&gt;(theta_final, energy_final)&lt;/code&gt; with the energy rounded to 6 decimals, by plain gradient descent.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;H = Z:   ground energy -1 at theta = 0
H = X:   ground energy -1 at theta = -pi/2 (mod 2pi)&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np

def energy(theta, H):
    ...

def grad_parameter_shift(theta, H):
    ...

def vqe(H, theta0, lr=0.3, steps=300):
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np

def _ry(theta):
    c, s = np.cos(theta / 2), np.sin(theta / 2)
    return np.array([[c, -s], [s, c]], dtype=complex)

def _state(theta):
    return _ry(theta) @ np.array([1, 0], dtype=complex)

def energy(theta, H):
    psi = _state(theta)
    return round(float(np.real(np.vdot(psi, np.asarray(H, dtype=complex) @ psi))), 6)

def _raw_energy(theta, H):
    psi = _state(theta)
    return float(np.real(np.vdot(psi, np.asarray(H, dtype=complex) @ psi)))

def grad_parameter_shift(theta, H):
    return round(0.5 * (_raw_energy(theta + np.pi / 2, H)
                        - _raw_energy(theta - np.pi / 2, H)), 6)

def vqe(H, theta0, lr=0.3, steps=300):
    th = float(theta0)
    for _ in range(steps):
        g = 0.5 * (_raw_energy(th + np.pi / 2, H) - _raw_energy(th - np.pi / 2, H))
        th = th - lr * g
    return (th, energy(th, H))
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Z = np.array([[1,0],[0,-1]], dtype=complex)
print(energy(0.0, Z), energy(np.pi, Z))
</text></testcode><stdin><text /></stdin><expected><text>1.0 -1.0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Z = np.array([[1,0],[0,-1]], dtype=complex)
print(grad_parameter_shift(np.pi/2, Z))
</text></testcode><stdin><text /></stdin><expected><text>-1.0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Z = np.array([[1,0],[0,-1]], dtype=complex)
print(vqe(Z, 0.3)[1])
</text></testcode><stdin><text /></stdin><expected><text>-1.0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
X = np.array([[0,1],[1,0]], dtype=complex)
print(vqe(X, 0.7)[1])
</text></testcode><stdin><text /></stdin><expected><text>-1.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
Z = np.array([[1,0],[0,-1]], dtype=complex)
X = np.array([[0,1],[1,0]], dtype=complex)
H = 0.6*Z + 0.8*X
print(vqe(H, 0.1)[1], round(float(np.linalg.eigvalsh(H)[0]), 6))
</text></testcode><stdin><text /></stdin><expected><text>-1.0 -1.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C10.Q3 — QAOA for MaxCut at depth p = 1</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C10.Q3 — Combinatorial optimisation, variationally&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;MaxCut on a graph \(G=(V,E)\) maximises \(C=\sum_{(i,j)\in E}\tfrac12(1-z_iz_j)\). Encode it as \(H_C=\sum_{(i,j)\in E}\tfrac12(I-Z_iZ_j)\), and use the mixer \(H_M=\sum_i X_i\).&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;maxcut_hamiltonian(n, edges)&lt;/code&gt;; &lt;code&gt;qaoa_state(n, edges, gamma, beta)&lt;/code&gt; returning \(e^{-i\beta H_M}e^{-i\gamma H_C}|+\rangle^{\otimes n}\); &lt;code&gt;qaoa_expectation(n, edges, gamma, beta)&lt;/code&gt; rounded to 6 decimals; and &lt;code&gt;best_cut(n, edges)&lt;/code&gt; returning the exact optimum by brute force over the \(2^{n}\) assignments.&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;triangle graph: best_cut = 2
square graph:   best_cut = 4&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np
from itertools import product

def maxcut_hamiltonian(n, edges):
    ...

def qaoa_state(n, edges, gamma, beta):
    ...

def qaoa_expectation(n, edges, gamma, beta):
    ...

def best_cut(n, edges):
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np
from itertools import product

I2 = np.eye(2, dtype=complex)
X = np.array([[0, 1], [1, 0]], dtype=complex)
Z = np.array([[1, 0], [0, -1]], dtype=complex)

def _embed(U, q, n):
    ops = [I2] * n
    ops[q] = U
    out = np.array([[1.0 + 0j]])
    for o in ops:
        out = np.kron(out, o)
    return out

def maxcut_hamiltonian(n, edges):
    H = np.zeros((2 ** n, 2 ** n), dtype=complex)
    for i, j in edges:
        H = H + 0.5 * (np.eye(2 ** n, dtype=complex) - _embed(Z, i, n) @ _embed(Z, j, n))
    return H

def _expm_diag(H, t):
    w, v = np.linalg.eigh((H + H.conj().T) / 2)
    return (v * np.exp(-1j * t * w)) @ v.conj().T

def qaoa_state(n, edges, gamma, beta):
    HC = maxcut_hamiltonian(n, edges)
    HM = np.zeros((2 ** n, 2 ** n), dtype=complex)
    for i in range(n):
        HM = HM + _embed(X, i, n)
    psi = np.full(2 ** n, 1 / np.sqrt(2 ** n), dtype=complex)
    psi = _expm_diag(HC, gamma) @ psi
    psi = _expm_diag(HM, beta) @ psi
    return psi

def qaoa_expectation(n, edges, gamma, beta):
    psi = qaoa_state(n, edges, gamma, beta)
    HC = maxcut_hamiltonian(n, edges)
    return round(float(np.real(np.vdot(psi, HC @ psi))), 6)

def best_cut(n, edges):
    best = 0
    for bits in product([0, 1], repeat=n):
        cut = sum(1 for i, j in edges if bits[i] != bits[j])
        best = max(best, cut)
    return best
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(best_cut(3, [(0,1),(1,2),(0,2)]), best_cut(4, [(0,1),(1,2),(2,3),(3,0)]))
</text></testcode><stdin><text /></stdin><expected><text>2 4
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
print(np.round(np.diag(maxcut_hamiltonian(2, [(0,1)])).real, 6).tolist())
</text></testcode><stdin><text /></stdin><expected><text>[0.0, 1.0, 1.0, 0.0]
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(qaoa_expectation(2, [(0,1)], 0.0, 0.0))
</text></testcode><stdin><text /></stdin><expected><text>0.5
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
n, edges = 2, [(0,1)]
best = max(qaoa_expectation(n, edges, g, b)
           for g in np.linspace(0, np.pi, 21) for b in np.linspace(0, np.pi, 21))
print(round(best, 4))
</text></testcode><stdin><text /></stdin><expected><text>0.9755
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
n, edges = 3, [(0,1),(1,2),(0,2)]
psi = qaoa_state(n, edges, 0.7, 0.4)
print(round(float(np.vdot(psi, psi).real), 6))
</text></testcode><stdin><text /></stdin><expected><text>1.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question><question type="coderunner"><name><text>C10.Q4 — Zero-noise extrapolation and shot budgets</text></name><questiontext format="html"><text>&lt;div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #c9a227;border-radius:8px;padding:18px 24px;margin-bottom:18px;"&gt;&lt;p style="color:#92400e;font-family:Georgia,'Times New Roman',serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 10px 0;"&gt;C10.Q4 — Error mitigation, and what it costs&lt;/p&gt;&lt;p style="color:#0f172a;font-family:Georgia,'Times New Roman',serif;font-size:15px;line-height:1.8;margin:0;"&gt;Zero-noise extrapolation measures an observable at several amplified noise levels \(\lambda_k\) and extrapolates to \(\lambda=0\). Richardson extrapolation is polynomial interpolation evaluated at zero.&lt;/p&gt;&lt;/div&gt;&lt;p style="font-family:Georgia,'Times New Roman',serif;font-size:15px;color:#0f172a;line-height:1.8;"&gt;Write &lt;code&gt;richardson(lams, vals)&lt;/code&gt; returning the value at \(\lambda=0\) of the interpolating polynomial through the given points (use the Lagrange formula, no &lt;code&gt;numpy.polyfit&lt;/code&gt; required but allowed), rounded to 6 decimals; and &lt;code&gt;shots_needed(variance, eps)&lt;/code&gt; returning \(\lceil \mathrm{Var}/\varepsilon^{2}\rceil\) as an &lt;code&gt;int&lt;/code&gt; — the sampling cost of estimating an expectation value to additive error \(\varepsilon\).&lt;/p&gt;&lt;div style="background:#fdf4ff;border:1px solid #e9d5ff;border-left:5px solid #9333ea;border-radius:8px;padding:14px 20px;margin:14px 0;"&gt;&lt;pre style="color:#3b0764;font-family:'Courier New',Consolas,monospace;font-size:13.5px;margin:0;background:none;border:none;"&gt;richardson([1,2,3], [0.9, 0.8, 0.7]) -&amp;gt; 1.0
shots_needed(1.0, 1e-3) -&amp;gt; 1000000&lt;/pre&gt;&lt;/div&gt;</text></questiontext><generalfeedback format="html"><text /></generalfeedback><defaultgrade>2.0000000</defaultgrade><penalty>0.0000000</penalty><hidden>0</hidden><idnumber /><coderunnertype>python3</coderunnertype><prototypetype>0</prototypetype><allornothing>0</allornothing><penaltyregime>0</penaltyregime><precheck>0</precheck><hidecheck>0</hidecheck><showsource>0</showsource><answerboxlines>18</answerboxlines><answerboxcolumns>100</answerboxcolumns><answerpreload>import numpy as np
from math import ceil

def richardson(lams, vals):
    ...

def shots_needed(variance, eps):
    ...
</answerpreload><globalextra /><useace /><resultcolumns /><template /><iscombinatortemplate /><allowmultiplestdins /><answer>import numpy as np
from math import ceil

def richardson(lams, vals):
    lams = [float(x) for x in lams]
    vals = [float(v) for v in vals]
    total = 0.0
    for i, (li, vi) in enumerate(zip(lams, vals)):
        term = vi
        for j, lj in enumerate(lams):
            if i != j:
                term *= (0.0 - lj) / (li - lj)
        total += term
    return round(total, 6)

def shots_needed(variance, eps):
    return int(ceil(float(variance) / float(eps) ** 2))
</answer><validateonsave>1</validateonsave><testsplitterre /><language /><acelang /><sandbox /><grader /><cputimelimitsecs /><memlimitmb /><sandboxparams /><templateparams /><hoisttemplateparams>1</hoisttemplateparams><extractcodefromjson>1</extractcodefromjson><templateparamslang>None</templateparamslang><templateparamsevalpertry>0</templateparamsevalpertry><templateparamsevald>{}</templateparamsevald><twigall>0</twigall><uiplugin /><uiparameters /><attachments>0</attachments><attachmentsrequired>0</attachmentsrequired><maxfilesize>10240</maxfilesize><filenamesregex /><filenamesexplain /><displayfeedback>1</displayfeedback><giveupallowed>0</giveupallowed><prototypeextra /><testcases><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(richardson([1,2,3], [0.9, 0.8, 0.7]))
</text></testcode><stdin><text /></stdin><expected><text>1.0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000"><testcode><text>print(shots_needed(1.0, 1e-3))
</text></testcode><stdin><text /></stdin><expected><text>1000000
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(richardson([1,3], [0.8, 0.4]))
</text></testcode><stdin><text /></stdin><expected><text>1.0
</text></expected><extra><text /></extra><display><text>SHOW</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>import numpy as np
lams = [1.0, 1.5, 2.0]
vals = [1 - 0.1*l + 0.02*l**2 for l in lams]
print(richardson(lams, vals))
</text></testcode><stdin><text /></stdin><expected><text>1.0
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase><testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000"><testcode><text>print(shots_needed(0.25, 0.01), shots_needed(1.0, 1.6e-3))
</text></testcode><stdin><text /></stdin><expected><text>2500 390625
</text></expected><extra><text /></extra><display><text>HIDE</text></display></testcase></testcases></question></quiz>