<?xml version="1.0" encoding="UTF-8"?>
<quiz>
  <question type="coderunner">
    <name><text>S13 Q1 — Verify an address checksum</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S13 Q1 — Address checksum</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">The last 4 bytes must equal the first 4 bytes of the double-SHA256 of everything before them.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write <strong><code>verify_checksum(data)</code></strong> where <code>data</code> is <code>payload + 4-byte checksum</code>. Return <code>True</code> if the checksum is correct. Use <code>hashlib</code>.</p>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[import hashlib

def verify_checksum(data):
    payload = data[:-4]
    checksum = data[-4:]
    # Complete: compare checksum to first 4 bytes of double SHA-256
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[import hashlib

def verify_checksum(data):
    payload = data[:-4]
    checksum = data[-4:]
    expected = hashlib.sha256(hashlib.sha256(payload).digest()).digest()[:4]
    return checksum == expected
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>import hashlib
payload = bytes([0x00]) + bytes(20)
chk = hashlib.sha256(hashlib.sha256(payload).digest()).digest()[:4]
print(verify_checksum(payload + chk))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>True
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>payload = bytes([0x00]) + bytes(20)
print(verify_checksum(payload + bytes([1, 2, 3, 4])))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>False
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>import hashlib
p = b"hello world payload"
chk = hashlib.sha256(hashlib.sha256(p).digest()).digest()[:4]
print(verify_checksum(p + chk))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>True
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name><text>S13 Q2 — Entropy in bits</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S13 Q2 — Entropy in bits</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">A secret of length L from a pool of P symbols has L &times; log2(P) bits of entropy.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write <strong><code>entropy_bits(pool, length)</code></strong> returning <code>length * log2(pool)</code>, rounded to 2 decimals. Use the <code>math</code> module.</p>
<div style="background:#fdf4ff;border-left:5px solid #9333ea;border-radius:8px;padding:12px 18px;margin:12px 0;">
<pre style="color:#3b0764;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;">entropy_bits(2048, 12) &rarr; 132.0
entropy_bits(26, 8)    &rarr; 37.6</pre>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[import math

def entropy_bits(pool, length):
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[import math

def entropy_bits(pool, length):
    return round(length * math.log2(pool), 2)
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(entropy_bits(2048, 12))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>132.0
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(entropy_bits(26, 8))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>37.6
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(entropy_bits(95, 10))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>65.7
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>

  <question type="coderunner">
    <name><text>S13 Q3 — Multisig backup tolerance</text></name>
    <questiontext format="html">
      <text><![CDATA[<div style="background:#f8fafc;border:1px solid #e2e8f0;border-left:5px solid #f7931a;border-radius:8px;padding:18px 22px;margin-bottom:16px;">
<p style="color:#9a3412;font-family:Georgia,serif;font-size:13px;font-weight:bold;letter-spacing:2px;text-transform:uppercase;margin:0 0 8px;">S13 Q3 — Backup tolerance</p>
<p style="color:#1e293b;font-family:Georgia,serif;font-size:15px;line-height:1.8;margin:0;">An m-of-n multisig still works after losing up to n &minus; m keys.</p>
</div>
<p style="font-family:Georgia,serif;font-size:15px;line-height:1.8;">Write <strong><code>keys_lost_tolerated(m, n)</code></strong> returning how many keys can be lost while still being able to spend.</p>
<div style="background:#fdf4ff;border-left:5px solid #9333ea;border-radius:8px;padding:12px 18px;margin:12px 0;">
<pre style="color:#3b0764;font-family:'Courier New',monospace;font-size:14px;margin:0;background:none;border:none;">keys_lost_tolerated(2, 3) &rarr; 1
keys_lost_tolerated(3, 5) &rarr; 2</pre>
</div>]]></text>
    </questiontext>
    <generalfeedback format="html"><text></text></generalfeedback>
    <defaultgrade>2.0000000</defaultgrade>
    <penalty>0.0000000</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <coderunnertype>python3</coderunnertype>
    <prototypetype>0</prototypetype>
    <allornothing>0</allornothing>
    <penaltyregime>0</penaltyregime>
    <precheck>0</precheck>
    <hidecheck>0</hidecheck>
    <showsource>0</showsource>
    <answerboxlines>18</answerboxlines>
    <answerboxcolumns>100</answerboxcolumns>
    <answerpreload><![CDATA[def keys_lost_tolerated(m, n):
    return ...
]]></answerpreload>
    <globalextra></globalextra>
    <useace></useace>
    <resultcolumns></resultcolumns>
    <template></template>
    <iscombinatortemplate></iscombinatortemplate>
    <allowmultiplestdins></allowmultiplestdins>
    <answer><![CDATA[def keys_lost_tolerated(m, n):
    return n - m
]]></answer>
    <validateonsave>1</validateonsave>
    <testsplitterre></testsplitterre>
    <language></language>
    <acelang></acelang>
    <sandbox></sandbox>
    <grader></grader>
    <cputimelimitsecs></cputimelimitsecs>
    <memlimitmb></memlimitmb>
    <sandboxparams></sandboxparams>
    <templateparams></templateparams>
    <hoisttemplateparams>1</hoisttemplateparams>
    <extractcodefromjson>1</extractcodefromjson>
    <templateparamslang>None</templateparamslang>
    <templateparamsevalpertry>0</templateparamsevalpertry>
    <templateparamsevald>{}</templateparamsevald>
    <twigall>0</twigall>
    <uiplugin></uiplugin>
    <uiparameters></uiparameters>
    <attachments>0</attachments>
    <attachmentsrequired>0</attachmentsrequired>
    <maxfilesize>10240</maxfilesize>
    <filenamesregex></filenamesregex>
    <filenamesexplain></filenamesexplain>
    <displayfeedback>1</displayfeedback>
    <giveupallowed>0</giveupallowed>
    <prototypeextra></prototypeextra>
    <testcases>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(keys_lost_tolerated(2, 3))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>1
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="1" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(keys_lost_tolerated(3, 5))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>2
</text></expected>
        <extra><text></text></extra>
        <display><text>SHOW</text></display>
      </testcase>
      <testcase testtype="0" useasexample="0" hiderestiffail="0" mark="1.0000000">
        <testcode><text>print(keys_lost_tolerated(3, 3))</text></testcode>
        <stdin><text></text></stdin>
        <expected><text>0
</text></expected>
        <extra><text></text></extra>
        <display><text>HIDE</text></display>
      </testcase>
    </testcases>
  </question>
</quiz>
